<img src='https://hammondm.github.io/hltlogo1.png' style="float:right">
Linguistics 531<br>
Fall 2024<br>
Jackson

## Things to remember about any homework assignment:

1. For this assignment, you will edit this jupyter notebook and turn it in. Do not turn in pdf files or separate `.py` files.
1. Late work is not accepted.
1. Given the way I grade, you should try to answer *every* question, even if you don't like your answer or have to guess.
1. You may *not* use `python` modules that we have not already used in class. (For grading, it needs to be able to run on my machine, and the way to do that is to limit yourself to the modules we've discussed and that are loaded into the Notebook.)
1. Don't use editors *other* than Jupyter Notebook to work on and submit your assignment, since they will mangle the autograding features: Google Colab, or even just editing the `.ipynb` file as a plain text file. Diagnosing and fixing that kind of problem takes a lot of my time, and that means less of my time to offer constructive feedback to you and to other students.
1. You may certainly talk to your classmates about the assignment, but everybody must turn in *their own* work. It is not acceptable to turn in work that is essentially the same as the work of classmates, or the work of someone on Stack Overflow, or the work of a generative AI model. Using someone else's code and simply changing variable or object names is *not* doing your own work.
1. All code must run. It doesn't have to be perfect, it may not do all that you want it to do, but it must run without error. Code that runs with errors will get no credit from the autograder.
1. Code must run in reasonable time. Assume that if it takes more than *5 minutes* to run (on your machine), that's too long.
1. Make sure to select `restart, run all cells` from the `kernel` menu when you're done and before you turn this in!

my name: Kathleen Costa

people I talked to about the assignment: N/A

# Homework #5

**This is due Tuesday, November 19, 2024 at noon (Arizona time).**

This assignment continues with the `NewB` corpus (downloadable [here](https://github.com/JerryWei03/NewB)).

imports:

```python
import re
from math import isclose

# Used in the cosimfreq() implementation from class:
import numpy as np
#  You're free to use it here in your functions, but it's not a requirement

# As with last week, this might make your life much easier
from collections import Counter
```

```python
!python --version
```

Output:
```
Python 3.12.7
```

**As before, this section is for autograding:**

What I need on my machine to properly grade this:

```python
# Path on my own machine, needed for GRADING
newbfile = '/home/ejackson1/Downloads/linguistics/NewB/train_orig.txt'

# ie, DON'T CHANGE THIS CELL, CHANGE THE ONE BELOW!
#  If you change *this* cell, the autograding is likely to break.
```

*In the editable cell below, enter the path on your own machine,* then uncomment that line so the notebook works on your machine.

**BEFORE YOU SUBMIT to D2L, remember to comment out *your* path again.**

```python
# YOUR path
newbfile = 'train_orig.txt'
```

We're going to build a tf-idf document index from the first 50,000 lines of the `train_orig.txt` file. We will *not* stem or remove stop words.

**1.** The first step is to build a term-document index and a document-term index. (8 points)

Follow the guidelines in the docstring below and in the following assert statements. As with previous assignments, you'll create a text processing function as part of this, and when you read in the file, retain only upper and lower case ASCII letters, numbers, and the percent sign. Remember that the document ID is *not* the same as the publication source code.

You should be able to re-use your code from previous assignments here. Note that what's new this week is (1) we've added a limit to how many lines of the input file we use, and (2) we're creating both a term-document index and a document-term index.

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
def makeIndices(filename,maxcount=50000):
    '''create document and term indices for the NewB corpus
    
    args:
        filename: location of the newB file
        maxcount: maximum number of lines
    
    returns:
        wordIndex: a dictionary from terms (as strings) to
            a list of tuples of the form (doc-id, count)
        docIndex: a list of tuples of the form
           (publication code (as an integer),
            text of document ie one sentence (as a list of strings),
            document vector (as a set of tuples: (word,count))
    '''
    # YOUR CODE HERE
    wordIndex = {}  
    docIndex = []   
    docId = 0
    
    with open(filename, 'r', encoding='utf-8') as f:
        lines = [next(f) for _ in range(maxcount)]
        
        for line in lines:
            parts = line.strip().split('\t')
            if len(parts) != 2:
                continue
                
            pubCode, text = int(parts[0]), parts[1]
            tokens = text_prep(text)
            
            doc_vector = {}
            for token in tokens:
                doc_vector[token] = doc_vector.get(token, 0) + 1
            
            doc_vector_set = set((term, count) for term, count in doc_vector.items())
            
            for term, count in doc_vector.items():
                if term not in wordIndex:
                    wordIndex[term] = []
                wordIndex[term].append((docId, count))
            
            docIndex.append((pubCode, tokens, doc_vector_set))
            docId += 1
    
    return wordIndex, docIndex
    
def text_prep(input):
    '''performs text normalization and tokenization on an input string
    
    Our process: anything that is not a letter (upper or lower case
        ASCII letters), digit (0-9), or the percent sign (%) is converted
        to space, and then terms are split on whitespace.
    
    args:
        input: a string of unprocessed text
    returns:
        output: a list of strings of normalized tokens
    '''
    # YOUR CODE HERE
    text = re.sub(r'[^a-zA-Z0-9%]', ' ', input)
    return [token.lower() for token in text.split() if token]
```

You've probably had enough testing of your `text_prep()` function that it should be working fine now. I won't include separate testing here, but bear in mind that if you are having trouble passing the tests below, it may be due to a problem in your `text_prep()` function. If so, the "practice" tests for that function from Homework 4 may be useful to you.

<!-- grade_id: q1t1, points: 1, grade: true, solution: false -->
```python
widx,didx = makeIndices(newbfile)

# test 1a, 1 pt
assert type(widx) == dict
```

<!-- grade_id: q1t2, points: 1, grade: true, solution: false -->
```python
# test 1b, 1 pt
# This one will be *very* sensitive to your text_prep() function
assert len(widx) == 32204
```

<!-- grade_id: q1t3, points: 1, grade: true, solution: false -->
```python
# test 1c, 1 pt
assert widx['apple'] == [(3984, 4),(10197, 1),(10604, 1),
 (27913, 1),(31235, 1),(33494, 1),(33511, 1),(46867, 1),
 (48326, 1),(48888, 1)]
```

<!-- grade_id: q1t4, points: 1, grade: true, solution: false -->
```python
# test 1d, 1 pt
assert type(didx) == list
```

<!-- grade_id: q1t5, points: 1, grade: true, solution: false -->
```python
# test 1e, 1 pt
assert len(didx) == 50000
```

<!-- grade_id: q1t6, points: 1, grade: true, solution: false -->
```python
# test 1f, 1 pt
assert didx[30861][0] == 1
```

<!-- grade_id: q1t7, points: 1, grade: true, solution: false -->
```python
# test 1g, 1 pt
assert didx[30861][1] == ['trump',  'holds',  'the',  'ground',  'lease',  'the',
                          'lease',  'for',  'the',  'land',  'on',  'which',  'the',
                          'building',  'stands']
```

<!-- grade_id: q1t8, points: 1, grade: true, solution: false -->
```python
# test 1h, 1 pt
assert didx[30861][2] == {('building', 1),  ('for', 1),  ('ground', 1),  ('holds', 1),
  ('land', 1),  ('lease', 2),  ('on', 1),  ('stands', 1),  ('the', 4),  ('trump', 1),
  ('which', 1)}
```

**2.** Use the term-document index to calculate idf values (an intermediate step on the way to ***tf*-idf** values) and return a dictionary from terms to idf values. (4 points)

Again, follow the guidance of the docstring and the following assert statements.

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
def makeIdfs(wordindex,maxcount=50000):
    '''calculate the idf value for each term in a word index
    
    args:
        wordindex: a term-document index as created by
            makeIndices()
        maxcount: total number of documents; by default,
            this is 50000
        
    returns:
        dictionary of each word in the collection
            and its idf value
    '''
    # YOUR CODE HERE
    idfs = {}
    
    for term, doc_list in wordindex.items():
        doc_count = len(doc_list)
        idfs[term] = np.log10(maxcount / doc_count)
        
    return idfs
```

<!-- grade_id: q2t1, points: 1, grade: true, solution: false -->
```python
idfx = makeIdfs(widx)

# test 2a, 1 pt
assert type(idfx) == dict
```

<!-- grade_id: q2t2, points: 1, grade: true, solution: false -->
```python
# test 2b, 1 pt
# There should be an entry here for every term in your term-document index
assert len(idfx) == 32204
```

<!-- grade_id: q2t3, points: 1, grade: true, solution: false -->
```python
# test 2c, 1 pt
# I get 4.698970004336019
assert isclose(idfx['canon'],4.69897,abs_tol=0.0001)
```

<!-- grade_id: q2t4, points: 1, grade: true, solution: false -->
```python
# test 2d, 1 pt
#I get 2.11690664142431
assert isclose(idfx['help'],2.11690,abs_tol=0.0001)
```

As we would expect, frequent terms (like function words) have very low idf values.

Because of the nature of this collection, certain other terms also have very low idf scores, even though they're not function words or common stop words, and would in other collections likely have ***higher*** idf scores. (So, note that idf scores will *really* depend on your data set!)

Here are the 20 terms with lowest idf scores; note that "donald" and "trump" are among them.

```python
sorted([(idfx[key], key) for key in idfx])[:20]
```

Output:
```
[(0.0, 'trump'),
 (0.22883893956150084, 'the'),
 (0.3805727251185078, 'to'),
 (0.4324090650491771, 'a'),
 (0.450457272115119, 'and'),
 (0.45859550737778554, 'of'),
 (0.5042059802251323, 'in'),
 (0.6440124083236123, 'that'),
 (0.6946916321358564, 'his'),
 (0.7188302265902644, 'on'),
 (0.7319839792180809, 'for'),
 (0.770471736212328, 'donald'),
 (0.7723189272471299, 'he'),
 (0.8082137524178006, 'said'),
 (0.8260565627156234, 'is'),
 (0.835290336460121, 'with'),
 (0.9376436819145622, 'has'),
 (0.9437051129237721, 'was'),
 (0.9510146974292889, 'as'),
 (1.0061230850587888, 'at')]
```

**3.** We now build a tf-idf index from the document-term index and the idf values. (6 points)

*HINT: This is really an index with precisely the same structure as our document-term index, but simply with different numerical values--floats, not just integer counts--in each document vector.*

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
def makeTfidf(idfs,docindex):
    '''creates a tfidf document-term index from idf scores
    and a document index
    
    args:
        idfs: idf scores (as a dictionary) as
            created by makeIdfs()
        docindex: a document index as created by
            makeIndices()
    returns:
        tf-idf index as a list where the list
            index corresponds to the docID
            and each entry is a triple:
                publication code
                document text
                set of pairs:
                    word
                    tfidf score
    '''
    # YOUR CODE HERE
    tfidf_index = []
    
    for doc_id, (pub_code, doc_text, term_counts) in enumerate(docindex):
        tfidf_scores = set()
        for term, count in term_counts:
            tfidf = count * idfs[term]
            tfidf_scores.add((term, tfidf))
        
        tfidf_index.append((pub_code, doc_text, tfidf_scores))
    
    return tfidf_index
```

<!-- grade_id: q3t1, points: 1, grade: true, solution: false -->
```python
tfidfidx = makeTfidf(idfx,didx)

# test 3a, 1 pt
assert type(tfidfidx) == list
```

<!-- grade_id: q3t2, points: 1, grade: true, solution: false -->
```python
# test 3b, 1 pt
assert len(tfidfidx) == 50000 and tfidfidx[49000][0] == 2
```

<!-- grade_id: q3t3, points: 1, grade: true, solution: false -->
```python
# test 3c, 1 pt
assert tfidfidx[49000][1] == ['donald', 'trump', 'didnt', 'tell', 'the', 'truth', '83', 'times', 'in', '1', 'day']
```

```python
# Here's a document we'll look at
print("Sentence:\n",tfidfidx[41584][1],'\n')
print("Raw counts:\n",sorted(didx[41584][2]),'\n')
print("tfidf:\n",sorted(tfidfidx[41584][2]))
```

Output:
```
Sentence:
 ['the', 'copying', 'of', 'the', 'letter', 'to', 'the', 'justice', 'department', 'attracted', 'wide', 'notice', 'in', 'washington', 's', 'close', 'knit', 'election', 'law', 'bar', 'as', 'did', 'the', 'claim', 'in', 'the', 'lawsuit', 'that', 'the', 'use', 'of', 'the', 'trump', 'foundation', 'to', 'benefit', 'the', 'trump', 'campaign', 'was', 'willful', 'and', 'knowing'] 

Raw counts:
 [('and', 1), ('as', 1), ('attracted', 1), ('bar', 1), ('benefit', 1), ('campaign', 1), ('claim', 1), ('close', 1), ('copying', 1), ('department', 1), ('did', 1), ('election', 1), ('foundation', 1), ('in', 2), ('justice', 1), ('knit', 1), ('knowing', 1), ('law', 1), ('lawsuit', 1), ('letter', 1), ('notice', 1), ('of', 2), ('s', 1), ('that', 1), ('the', 8), ('to', 2), ('trump', 2), ('use', 1), ('was', 1), ('washington', 1), ('wide', 1), ('willful', 1)] 

tfidf:
 [('and', 0.450457272115119), ('as', 0.9510146974292889), ('attracted', 3.3767507096020997), ('bar', 3.0087739243075053), ('benefit', 2.703
```

<!-- grade_id: q3t4, points: 1, grade: true, solution: false -->
```python
# I'm going to put the tf-idf representation of the document into a Counter so it's
#  easier to pull values out for a specific term (using the Counter like a dict)
res = Counter(dict(tfidfidx[41584][2]))

# test 3d, 1 pt
# I get 0 and 4.221848749616356
assert isclose(res['trump'],0.,abs_tol=0.0001) and isclose(res['willful'],4.2218,abs_tol=0.0001)
# I get 0.9171910147555711 and 1.8307115164920067
assert isclose(res['of'],0.91719,abs_tol=0.0001) and isclose(res['the'],1.83071,abs_tol=0.0001)
```

```python
# Here's another document that we'll look at
print(didx[43272][1])
```

Output:
```
['journalism', 'somehow', 'managed', 'to', 'get', 'the', 'trump', 'phenomenon', 'wrong', 'wrong', 'wrong', 'while', 'writing', 'more', 'and', 'more', 'about', 'trump', 'trump', 'trump']
```

<!-- grade_id: q3t5, points: 1, grade: true, solution: false -->
```python
res = Counter(dict(tfidfidx[43272][2]))

# test 3e, 1 pt
# I get 2.8459360287479374 and 7.517536217944672
assert isclose(res['more'],2.8459,abs_tol=0.0001) and isclose(res['wrong'],7.5175,abs_tol=0.0001)
```

```python
# Here's another document that we'll look at
print(didx[7328][1])
```

Output:
```
['but', 'when', 'you', 'have', 'a', 'fight', 'you', 'get', 'sinatra', 'you', 'get', 'donald', 'trump', 'you', 'get', 'lee', 'iacocca', 'and', 'you', 'get', 'dr']
```

<!-- grade_id: q3t6, points: 1, grade: true, solution: false -->
```python
res = Counter(dict(tfidfidx[7328][2]))

# test 3f, 1 pt
# I get 8.040034638701437 and 7.415487857287048
assert isclose(res['you'],8.0400,abs_tol=0.0001) and isclose(res['get'],7.4155,abs_tol=0.0001)
```

We'll need the `cosimfreq()` function again:

```python
#cosine similarity wrt/frequencies or tfidfs
def cdot(count1, count2):
    """dot product for two vectors as Counters"""
    return sum(count1[word]*count2[word] for word in count1)

def clength(count1):
    """length of a vector as Counter"""
    return np.sqrt(sum(count1[word]**2 for word in count1))

def cosimfreq(d1,d2):
    # gotta get 'em from lists of tuples into Counters
    d1, d2 = Counter(dict(d1)), Counter(dict(d2))
    d1len, d2len = clength(d1), clength(d2)
    denom = d1len * d2len
    if denom == 0: return 0
    return float(cdot(d1, d2))/float(denom)
```

**4.** Write a search function for the tf-idf index from question 3 that uses only cosine similarity (as given just above). (5 points)

The function should have this argument structure:

```python
search(query, index, idfs)
```

The function should return a rank-ordered list of the 10 **best** matches.

*Hint: Note that the idf scores, as well as the tf-idf index and query, are being passed to this search function. Consider the last several homework assignments: often you made a change to the structure of an index, which meant that you needed to make the same change to the distance functions and to the search function. What did you have to update in the search functions before?*

<!-- grade_id: q4a, points: 0, grade: false, solution: true -->
```python
def search(query,index,idfs):
    '''search a tf-idf index for documents that are most similar
       to a query
       (This time we won't offer a choice of distance function,
       but will just use cosine similarity.)
    
    args:
        query: search query (as a string)
        idx: tf-idf index (as a list of tuples)
        idfs: idf scores (as a dictionary)
    returns:
        10 best matches as a list of tuples:
            cosine similarity score
            document index
    '''
    # YOUR CODE HERE
    query_terms = re.sub(r'[^a-zA-Z0-9%]', ' ', query.lower()).split()
    
    query_counts = Counter(query_terms)
    query_vector = {}
    for term in query_terms:
        if term in idfs:
            query_vector[term] = query_counts[term] * idfs[term]
    
    results = []
    for doc_id, (_, _, doc_vector) in enumerate(index):
        doc_dict = dict(doc_vector)
        
        numerator = sum(query_vector.get(term, 0.0) * doc_dict.get(term, 0.0) 
                       for term in set(query_vector) | set(doc_dict))
        
        query_magnitude = np.sqrt(sum(v * v for v in query_vector.values()))
        doc_magnitude = np.sqrt(sum(v * v for v in doc_dict.values()))
        
        if query_magnitude and doc_magnitude:
            similarity = float(numerator / (query_magnitude * doc_magnitude))
            results.append((similarity, doc_id))
        else:
            results.append((0.0, doc_id))
    
    return sorted(results, reverse=True)[:10]
```

<!-- grade_id: q4t1, points: 1, grade: true, solution: false -->
```python
res1 = search("melania is trump's wife",tfidfidx,idfx)

# test 4a, 1 pt
# Make sure you're getting the structure right
assert type(res1) == list and type(res1[0]) == tuple and type(res1[0][0]) == float and type(res1[0][1]) == int
```

<!-- grade_id: q4t2, points: 1, grade: true, solution: false -->
```python
# test 4b, 1 pt
# Make sure you're getting the right results returned
assert res1[0][1] == 29285 and res1[1][1] == 41827
```

```python
# As a sanity check, here are your results:
[(result,' '.join(didx[result[1]][1])) for result in res1]
```

Output:
```
[((0.849969028092293, 29285), 'trump and his wife melania'),
 ((0.6215127702787681, 41827), 'trump and his wife melania hosted mr'),
 ((0.5183980745617279, 46573),
  'full text here is a look at the life of melania trump wife of president donald trump'),
 ((0.5083112011174594, 21175),
  'trump is 24 years older than his wife melania'),
 ((0.5043656908492963, 10060),
  'from left trump his son barron trump and wife melania trump'),
 ((0.5013248622362927, 32775),
  'trump and his wife melania who was pregnant at the time'),
 ((0.4673037889163044, 15962),
  'trump was accompanied by his wife melania who introduced him'),
 ((0.45156314335095366, 21510),
  'he then introduced his wife melania who said trump will work for all americans'),
 ((0.45156314335095366, 3499),
  'he then introduced his wife melania who said trump will work for all americans'),
 ((0.4364630422890691, 47305),
  'full text here is a look at the life of melania trump wife of 45th us president donald trump')]
```

<!-- grade_id: q4t3, points: 1, grade: true, solution: false -->
```python
# test 4c, 1 pt
# Make sure you're getting the right similarity values, too
# For the top two, I get 0.849969028092293 and 0.6215127702787681
assert isclose(res1[0][0],0.84997,abs_tol=0.0001) and isclose(res1[1][0],0.62151,abs_tol=0.0001)
```

<!-- grade_id: q4t4, points: 1, grade: true, solution: false -->
```python
res2 = search('trump lives in the white house',tfidfidx,idfx)

# test 4d, 1 pt
assert res2[0][1] == 22752 and res2[4][1] == 45145
```

```python
# Sanity check
[(result,' '.join(didx[result[1]][1])) for result in res2]
```

Output:
```
[((0.6319877962507022, 22752), 'trump in the white house'),
 ((0.6158608687047019, 11090), 'the trump white house'),
 ((0.6067586288920347, 31782), 'trump to the white house'),
 ((0.5931508697856988, 37704), 'trump is in the white house'),
 ((0.49816997165817267, 45145), 'trump and the white house said mr'),
 ((0.4539447710242628, 13798), 'full text the trump white house'),
 ((0.42225789139265635, 15188), 'he lives on the phone walker said of trump'),
 ((0.40932614931370964, 2490), 'we can go on with our lives ivana trump said'),
 ((0.4027002229591978, 6749), 'one of the characters lives in trump towers'),
 ((0.39402213680464887, 26481), 'trump was in the white house for a reason')]
```

<!-- grade_id: q4t5, points: 1, grade: true, solution: false -->
```python
# test 4e, 1 pt
# I get 0.3940221368046489 for the tenth result
assert isclose(res2[9][0],0.3940, abs_tol=0.0001)
```

**5.** Calculate the MAP score for your search function based on `res1` and `res2` above *by hand* (that is, step by step, typed below) and *show your work*. (2 points)

Note that this requires that you look up the sentence for each returned result and judge whether it was relevant to the query. Your choice of relevant or not isn't graded, so don't invest too much thought in what *ought* to be relevant or not; choose a judgment for each returned result, and then you'll need to properly calculate the MAP score based on those choices.

First, here is some text formatting code in Markdown for you to use. Copy the source of this cell into the answer cell (since this cell will be read-only), then give your relevance scores in this table. Code each result as "relevant" (Y) or "not relevant" (N) for its query&mdash;meaning, is this result a good result given that particular query.

As an example of what it means for one of our results to be "relevant" for the queries here, consider our `res2` query, "trump lives in the white house". Many of the results involve the terms "trump" and "white house", and to me they do seem like reasonable results for that query, but note that result 15188 is "he lives on the phone walker said of trump". To me, that seems like **not** a relevant result for "trump lives in the white house". However, the final decision is up to you, and I'm the one who has to make sure that the math is correct given *your* relevancy results.

&nbsp; | `res1` | `res2`
-------|--------|-------
0
1
2
3
4
5
6
7
8
9

Based on your table, once you generate it in the cell below, please show the step-by-step calculation of the MAP score for this search system. You may find the following references helpful if you need to properly format any of your mathematical equations in LaTeX and Markdown:

https://jupyterbook.org/en/stable/content/math.html

https://ashki23.github.io/markdown-latex.html#latex

<!-- grade_id: q5a, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

&nbsp; | `res1` | `res2`
-------|--------|-------
0      |   Y    |   Y
1      |   Y    |   Y
2      |   Y    |   Y
3      |   Y    |   Y
4      |   Y    |   N
5      |   Y    |   N
6      |   Y    |   N
7      |   Y    |   N
8      |   Y    |   N
9      |   Y    |   Y

For res1, since all the results were relevant, we will calculate the precision of the relevant lines: 
$$
\mathrm{AP}(q) = \frac{1}{m} \sum_{k=1}^{m} \mathrm{Precision}(R_k)
$$
Here, the number of relevant documents I found was 10, and that is the overall number of documents present so the calculation for res1
would be:
   - Rank 1: $$ \text{Precision} = \frac{1}{1} = 1.0  $$ 
   - Rank 2: $$ \text{Precision} = \frac{2}{2} = 1.0 $$  
   - Rank 3: $$ \text{Precision} = \frac{3}{3} = 1.0 $$  
   - Rank 4: $$ \text{Precision} = \frac{4}{4} = 1.0 $$  
   - Rank 5: $$ \text{Precision} = \frac{5}{5} = 1.0 $$  
   - Rank 6: $$ \text{Precision} = \frac{6}{6} = 1.0 $$  
   - Rank 7: $$ \text{Precision} = \frac{7}{7} = 1.0 $$  
   - Rank 8: $$ \text{Precision} = \frac{8}{8} = 1.0 $$  
   - Rank 9: $$ \text{Precision} = \frac{9}{9} = 1.0 $$  
   - Rank 10: $$ \text{Precision} = \frac{10}{10} = 1.0 $$  


$$
\text{AP} = \frac{1.0+1.0+1.0+1.0+1.0+1.0+1.0+1.0+1.0+1.0=}{10} = 1.0
$$

The reason I chose all the documents as relevant here is because they all mention Trump, Melania, and the word "wife" to describe Melania. 


Using the same equation for res2, the calculation of the precision of relevant lines is:
   - Rank 1: $$ \text{Precision} = \frac{1}{1} = 1.0 $$  
   - Rank 2: $$ \text{Precision} = \frac{2}{2} = 1.0 $$  
   - Rank 3: $$ \text{Precision} = \frac{3}{3} = 1.0 $$  
   - Rank 4: $$ \text{Precision} = \frac{4}{4} = 1.0 $$  
   - Rank 10: $$ \text{Precision} = \frac{5}{10} = 0.5 $$  


$$
\text{AP} = \frac{1.0+1.0+1.0+1.0+0.5=}{5} = 0.9 
$$

The reason I chose these specific documents as relevant is because while some documents do mention Trump in the White House, some of the
results were irrelevant and discussed other members of the Trump family or other locations that the Trumps might frequent.

To calculate the overall MAP, we will need to take an average of the MAPs for both results. That would look like:
$$
\text{MAP} = \frac{\text{AP(res1)} + \text{AP(res2)}}{2} = \frac{1.0 + 0.9}{2} = 0.95
$$

Overall, the precision of this model is 0.95. 